# <span style="color:teal;">CIS 211 Project 2:  Cryptography</span>

##### Due 11:00 P.M. April 14, 2017

##### Reading:  M&R 3.1 -- 3.3, 3.5; &nbsp; 8.1 -- 8.2, 8.4.1 -- 8.4.4

This week we're going to use simple encryption and decryption algorithms as a way of exploring strings and containers (lists, dictionaries, and sets) in Python.

Encryption is described in Chapter 3 of the textbook.  You should understand how the **substitution cipher** works -- you don't need to understand the details of any of the Python functions that implement the method, but you should know what it means to encode a piece of text with this kind of cipher.

Writing a function that decrypts a piece of text (assuming it really was encrypted with a substitution cipher) is not too hard, compared to other forms of encryption, but still more than we want to do in a one-week project.  Instead of developing a complete algorithm that will decrypt a piece of text we'll write a few small "helper functions" that might be useful as part of a larger application.


###  <span style="color:teal">Imports</span>

```python
from string import punctuation, ascii_uppercase, ascii_lowercase
```

###  <span style="color:teal">Test Data</span>

```python
seuss = '\nUEK XZIV, MOU XZIV,\nRKP XZIV, QHGK XZIV.\n'
print(seuss)
```

Output:
```
UEK XZIV, MOU XZIV,
RKP XZIV, QHGK XZIV.
```

###  <span style="color:teal">Project 1: &nbsp; Ciphertext Words (20 points)</span>

Fill in the body of the function named `cipher_words`.  It will be passed a string containing a "cryptoquote", and it should return the set of words in the quote.  All punctuation should be stripped from the ends of words (but not in the middle).

Example:
```
>>> cipher_words(suess)
{'MOU', 'QHGK', 'RKP', 'UEK', 'XZIV'}
```

**Style (2 points):** This function can be written without a `for` loop, using set comprehension.

##### <span style="color:red">Code</span>

<!-- grade_id: cipher_words, points: 10, grade: true, solution: true -->
```python
def cipher_words(text):
    '''
    Takes a string and returns a list of all unique words in the string.
    
    :param str text: String of a cipher
    :return: A set of all the unique words in the cipher
    :rtype: set
    '''
    
    thing = set(text.strip().split())
    return {part.strip(punctuation) for part in thing}
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: cipher_words_test, points: 5, grade: true, solution: false -->
```python
# pass the example text to the function, make sure 5 words are returned
words = cipher_words(seuss)
assert len(words) == 5

# look up each word in the result set
for w in ['MOU', 'QHGK', 'RKP', 'UEK', 'XZIV']:
    assert w in words
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: cipher_words_doc, points: 5, grade: true, solution: true -->
Takes a string and strips it into a set with all of the unique words, but with punctuation still intact (thing). The function then returns a set of all unique words by taking the words in thing, stripping them of all punctuation, and then placing them into a set. Any words that are the same, but were determined to be different because of the punctuation are taken care of in this step.

Returns a set of all the unique words. 

###  <span style="color:teal">Project 2: &nbsp; Count Ciphertext Letters (20 points)</span>

Fill in the body of the `count_unsolved` function.  It will be passed a single word, and it should return the number of unique ciphertext (upper case) letters in the word.

Examples:
```
>>> count_unsolved('XZIV')
4

>>> count_unsolved('QTETET')
3
```
Note how there are 6 letters in the second example, but some are repeated.  There are only 3 different letters:  Q, E, and T.

##### <span style="color:red">Code</span>

<!-- grade_id: count_unsolved, points: 10, grade: true, solution: true -->
```python
def count_unsolved(word):
    '''
    Takes a single word and returns the number of unsolved letters.
    
    :param str word: Word that has some combination of solved and unsolved letters.
    :return: The number of letters that remain unsolved
    :rtype: int
    '''
    unique = ""
    count = 0
    for char in word:
        if char in ascii_uppercase: 
            if char not in unique:
                unique += char
                count += 1
    return count
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: count_unsolved_test_1, points: 1, grade: true, solution: false -->
```python
# each letter in this string is unique
assert count_unsolved('XZIV') == 4
```

<!-- grade_id: count_unsolved_test_2, points: 1, grade: true, solution: false -->
```python
# there are only 3 unique letters in this string
assert count_unsolved('QTETET') == 3
```

<!-- grade_id: count_unsolved_test_3, points: 1, grade: true, solution: false -->
```python
# repeat the previous test after "solving" one of the letters
assert count_unsolved('QaEaEa') == 2
```

<!-- grade_id: count_unsolved_test_4, points: 1, grade: true, solution: false -->
```python
# a word can be completely solved (no remaining ciphertext letters)
assert count_unsolved('banana') == 0
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: count_unsolved_doc, points: 6, grade: true, solution: true -->
For each character, the function ascertains whether the character is an uppercase letter. If it is an uppercase letter, it checks to see if the character is a letter that has occurred earlier in the word (these letters are stored in the string unique). If it is uppercase and it is not in unique, count goes up by one and the letter is added to unique. 

After going through all the characters, the function returns the number of letters that remain unsolved.

###  <span style="color:teal">Project 3: &nbsp; Fewest Unsolved Letters (15 points)</span>

One strategy for choosing words to work on is to find the word that has the fewest letters remaining to be solved.  Fill in the body of the `fewest_unsolved` function so it iterates over a set of words and returns a tuple `(n,w)` where `w` is the word with the fewest remaining undecrypted letters and `n` is the number of undecrypted letters.

For example, suppose the quote from earlier has been partially decrypted and the set of words is now
```
{'MOo', 'QHue', 'ReP', 'XiIV', 'oEe'}
```
Passing this set to `fewest_unsolved` should return `oEe` because it has one uppercase letter and the others all have two:
```
```

**Note:** If there is a tie (more than one word has the same minimum number of undecrypted letters) just return one at random (_e.g._ the first one found).

##### <span style="color:red">Code</span>

<!-- grade_id: fewest_unsolved, points: 6, grade: true, solution: true -->
```python
def fewest_unsolved(group):
    '''
    Takes a set of words and returns the word that has the fewest number unsolved. 
    
    :param set group: A set of partially ciphered words
    :return: The word that has the fewest number of words unsolved
    :rtype: str
    '''
    fewest_str = 'QQQQQQQQQQQQQQQQQQQQQQQQQQQQQQQQQQQ'
    fewest_no = 999999999999999
    for item in group:
        if count_unsolved(item) < fewest_no:
            fewest_str = item
            fewest_no = count_unsolved(item)

    return fewest_str
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: fewest_unsolved_test1, points: 2, grade: true, solution: false -->
```python
# this set has a word with one ciphertext letter and three others with two ciphertext letters
test_set = {'oEe', 'XiIV', 'MOo', 'QHue', 'ReP'}

# the first call should find the word with one ciphertext letter
w = fewest_unsolved(test_set)
assert w == 'oEe'
```

<!-- grade_id: fewest_unsolved_test_2, points: 2, grade: true, solution: false -->
```python
# remove that word from the set, the next call should find a word with two ciphertext letters
test_set.remove(w)
w = fewest_unsolved(test_set)
assert count_unsolved(w) == 2
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: fewest_unsolved_doc, points: 5, grade: true, solution: true -->
Takes a set of partially ciphered words and returns the word with the fewest unsolved letters. The values of fewest_str and fewest_no are set intentionally high so that the first word in the set will always replace these values. For each word in the set, the function calls count_unsolved. If the value of count_unsolved is strictly less than the current lowest number unsolved (fewest_no), then fewest_str is replaced by the word in question, and fewest_no is replaced with the value of count_unsolved. 

Returns the word with the fewest unsolved letters. 

###  <span style="color:teal">Project 4: &nbsp; Letter Pairs (30 points)</span>

An effective strategy for decrypting English text is to consider letters two at a time (see Section 8.4.3 of the textbook for an explanation of why this approach is helpful).

For this part of the project, fill in the definition of the function named `make_letter_pairs`.  It will be passed a string, and it should return a matrix $m$, such that $m_{i,j}$ is the frequency with which letters $i$ and $j$ were next to each other in the string.

One way to create a matrix in Python is to use a list of lists, but for this project we'll make a dictionary of dictionaries.  If we save the result of the function call in a variable named `m`, we can look up a frequency using indexing.  For example, we can look up the frequency of `'i'` followed by `'e'` by evaluating
```
>>> m['i']['e']
```

Here is an example:
```
>>> m = make_letter_pairs('how now, brown cow?')
>>> m
{'b': {'r': 0.1},
 'c': {'o': 0.1},
 'h': {'o': 0.1},
 'n': {'o': 0.1},
 'o': {'w': 0.4},
 'r': {'o': 0.1},
 'w': {'n': 0.1}}
```

The 10 places in this text where there are two letters next to each other are `ho`, `ow`, `no`, `ow`, `br`, `ro`, `ow`, `wn`, `co`, and `ow`.  The combination of `b` followed by `r` occurs only once, so its frequency is 1 / 10 = 0.1:
```
>>> m['b']['r']
0.1
```

The pair `ow` occurs four times, so its frequency is 4 / 10 = 0.4:
```
>>> m['o']['w']
0.4
```

**Note**  There is a similar function, named `neighborCount`, defined in the book. You can use this function to get some ideas for how to write your code, but there are some important differences you should be aware of:
* the function in the book creates counts, but our function creates frequencies
* the function in the book counts in both directions, _e.g._ when it sees the pair `ow` it adds 1 to the count of `ow` and one to the count for `wo`; we count one direction only

##### <span style="color:red">Code</span>

<!-- grade_id: letter_pairs, points: 15, grade: true, solution: true -->
```python
def make_letter_pairs(text):
    '''
    Returns a dictionary with the frequency of letter pairs in a given text.
    
    :param str text: Text that is being investigated for the letter pairs.
    :return: A dictionary with letter pairs and their frequency in the given text.
    :rtype: dict
    '''
    
    pair_dict = {}
    text = text.lower()
    letters = 0
    # Finding the number of pairs of letters in the text
    for j in range(len(text)):
        if text[j] in ascii_lowercase:
            if text[j+1] in ascii_lowercase:
                letters += 1
    
    # Placing the letter pairs into the dictionary
    for i in range(len(text) - 1):
        if text[i] in ascii_lowercase:
            if text[i] in pair_dict:
                if text[i+1] in pair_dict[text[i]] and text[i+1] in ascii_lowercase:
                    pair_dict[text[i]][text[i+1]] += 1 / letters
                elif text[i+1] not in pair_dict[text[i]] and text[i+1] in ascii_lowercase:
                    pair_dict[text[i]][text[i+1]] = 1 / letters
            elif text[i] not in pair_dict and text[i+1] in ascii_lowercase:
                pair_dict[text[i]] = {text[i+1]: 1 / letters}

    return pair_dict
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: letter_pairs_test_1, points: 4, grade: true, solution: false -->
```python
# the first test uses the 'brown cow' example
m = make_letter_pairs('how now, brown cow?')

# there are 7 different letters in the string
assert len(m) == 7

# check two of the frequencies
assert round(m['o']['w'],1) == 0.4
assert round(m['r']['o'],1) == 0.1
```

<!-- grade_id: letter_pairs_test_2, points: 2, grade: true, solution: false -->
```python
# the remaining tests are based on frequencies using a large body of text (War of the Worlds, by H. G. Wells)
with open('wells.txt') as f:
    pp = make_letter_pairs(f.read())

# look up some common letter combinations ("pp" stands for "pair probability")
assert round(pp['e']['e'], 3) == 0.006
assert round(pp['s']['s'], 3) == 0.004
```

<!-- grade_id: letter_pairs_test_3, points: 2, grade: true, solution: false -->
```python
# look up some rare combinations (in this book, at least)
assert round(pp['r']['z']) < 0.000001
assert round(pp['y']['n']) < 0.000001
```

<!-- grade_id: letter_pairs_test_4, points: 2, grade: true, solution: false -->
```python
# letter combinations 'jx' and 'kk' never occur in this book
assert 'x' not in pp['j']
assert 'k' not in pp['k']
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: letter_pairs_doc, points: 5, grade: true, solution: true -->
```python
'''
The function initially determines how many letter pairs there are in the text by determining if the ith and i + 1th character of the text are both lowercase letters (The entirety of the text has been lowered via text.lower()). Then for each character, it determines whether (1) if the current character is a lowercase letter, (2) if the current character is in the pair dictionary, (3) if the subsequent character is a lowercase letter, (4) if the subsequent character in the dictionary for the current character that is nested in the pair dictionary. 

If (1) is True, the function moves on to step 2.
If (1) is False, the function moves on to the next character.

If (2) is True, then the function moves on to step 3.
If (2) is False, then the function moves on to step 3.

If (3) is True, then the function moves on to step 4.
If (3) is False, then the function moves on to the next character.

If (4) is True, then the function adds (1 / (total number of letter pairs)) to the value of the dictionary key that corresponds to the letter pair.
If (4) is False, the dictionary for the current letter is updated to add a key of the subsequent letter with a value of (1 / total number of letter pairs). 

The function returns the pair dictionary. 
'''
```

###  <span style="color:teal">Project 5: `apply` (15 points)</span>

The final part of the project is to write a function named `apply` that will use a decryption key (a dictionary mapping ciphertext letters to plaintext) to update a partially decrypted string.

For example, this key replaces X with `'f'`, Z with `'i'`, I with `'s'`, and V with `'h'`:
```
>>> k = {'X':'f', 'Z':'i', 'I':'s', 'V':'h'}
```

If we apply the key to the example ciphertext this is the result:
```
>>> apply(k, 'UEK XZIV, MOU XZIV, RKP XZIV, QHGK XZIV.')
'UEK fish, MOU fish, RKP fish, QHGK fish.'
```

##### <span style="color:red">Code</span>

<!-- grade_id: apply, points: 6, grade: true, solution: true -->
```python
def apply(key, quote):
    '''
    Replaces decrypted letters in a ciphertext with their known values.
    
    :param set key: The key to the ciphertext.
    :param str quote: The quote that is being attempted to be deciphered.
    :return: A string with the known letters replaced within the string.
    :rtype: str
    '''
    
    final_string = ''
    i = 0
    for char in list(quote):
        i += 1
        for thing in key:
            if char == thing:
                final_string += key[thing]
        if len(final_string) != i:
            final_string += char
    return final_string
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: apply_test, points: 4, grade: true, solution: false -->
```python
# make a decryption key to use in the tests
k = {'X':'f', 'Z':'i', 'I':'s', 'V':'h'}

# apply the key 
assert apply(k,'XZIV') == 'fish'
assert apply(k,'UDKLCXZIV') == 'UDKLCfish'
assert apply(k,'XMIVZQP') == 'fMshiQP'
assert apply(k,'JSBWC') == 'JSBWC'
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: apply_doc, points: 5, grade: true, solution: true -->
For each character in quote, the function checks each key in the key dictionary. If they match, then the value of the specific key is placed into the string final_string. If after going through all of the keys there is no match, the original letter is placed into final_string. 

After going through all of the letters in quote, the function returns the final string.