# <span style="color:teal;">CIS 211 Project 1:  Hello, World</span>

##### Due 11:00 P.M. April 7, 2017

##### Reading:  none

## Instructions

This goal for this project is to work through the process for completing projects in CIS 211 this term:
* download a project notebook from Canvas, save it as `hello.ipynb` 
* edit the file with Jupyter to add code and documentation
* test your code
* submit the completed notebook using Canvas

The notebooks you download from Canvas will have specially tagged cells where you should write your code and documentation.  These cells will be read by the **autograder**, a program that automatically loads your program and runs a set of unit tests.  

The specially tagged cells are immediately below red headlines that say **Code** or **Unit Tests** or **Documentation**.  

**<span style="color:red">Important!! &nbsp; Do not delete or move the specially tagged cells.</span>** 

##  <span style="color:teal">The `hello` Function</span>

Write a function named `hello` that will take a language name passed as a string and return a "hello, world" greeting in that language.

Examples:
<pre>
>>> hello('english')
'Hello, World!'

>>> hello('french')
'Bonjour, tout le monde!'

>>> hello('arabic')
'مرحبا يا عالم'
</pre>

Your function should be able to return a greeting in at least two languages -- English and French -- but you can have as many other languages as you would like.

The strings to return for English and French are shown in the tests below.  Make sure you return these strings exactly as shown so your program passes the unit tests.

If the language passed as an argument is not one of the languages you implemented return None.
<pre>
>>> hello('valyrian')
</pre>

##### Style Points:

(1) Try to write your program without using a bunch of `if-elif` cases. This is <span style="color:red"> not</span> a good way to structure the code in the body of your `hello` function:
<pre>
if lang == 'english':                   # bad style
    reply = 'Hello, World!'
elif lang == 'french':                  # ick
    reply = 'Bonjour, tout le monde!'
elif:                                   # ugh
    ...
</pre>

(2) How do you want to handle situations where the language name is capitalized?  For full credit your function should work when called either as `hello("english")` or `hello("English")`.

##### <span style="color:red">Code</span>

<!-- grade_id: hello_code, points: 50, grade: true, solution: true -->
```python
# YOUR CODE HERE
def hello(language):
    language = language.lower()
    lang_dict = {'english': "Hello, World!",
                'french': 'Bonjour, tout le monde!',
                'spanish': '¡Hola, mundo!',
                'german': 'Hallo welt!',
                'portuguese': 'Olá Mundo!',
                'dutch': 'Hallo Wereld!',
                'italian': 'Ciao mondo!',
                'mime': '',
                'norwegian': 'Hei Verden!'}
    
    if language in lang_dict:
        return lang_dict[language]
    else:
        return None
```

##### <span style="color:red">Auto-Grader Unit Tests</span>

<!-- grade_id: english, points: 10, grade: true, solution: false -->
```python
assert hello('english') == 'Hello, World!'
```

<!-- grade_id: french, points: 10, grade: true, solution: false -->
```python
assert hello('french') == 'Bonjour, tout le monde!'
```

<!-- grade_id: klingon, points: 10, grade: true, solution: false -->
```python
assert hello('klingon') is None
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: hello_doc, points: 20, grade: true, solution: true -->
Takes the name of a language as a string and returns "Hello, World!" in whichever language is inputted. if that language is not accounted for, the function returns None.

(language) -> Hello, World in that language if that language is accounted for. 

(language) -> None if that language is not accounted for. 