# <span style="color:teal;">CIS 211 Project 5:  Building GUIs with Tk</span>

##### Due 11:00 P.M. Friday May 5

Reading:
* [TkDocs - Tk Tutorial](http://www.tkdocs.com/tutorial/index.html)
* [Tkinter Tutorial](http://www.python-course.eu/python_tkinter.php) from `python-course.eu`

### <span style="color:teal">Setup</span>

Execute the following code cell to tell IPython you want to create windows with Tk and to import the `tkinter` library.

```python
%gui tk
import tkinter as tk
```

<span style="color:red">**Important:**</span> &nbsp; The autograder tests assume `tkinter` has been imported as shown above.  You should follow the same convention, and refer to widget classes as `tk.Label`, `tk.Entry`, and so on.

<span style="color:red">**Important:**</span> &nbsp; The code cells where you define your GUIs contain assignment statements that create global variables.  For example, the first GUI is assigned to a variable named `hello_app`.  **Do not change these assignment statements**.  The autograder test cells will use these variables to access parts of your GUIs.

### <span style="color:teal">1. &nbsp; Hello, World (10 points)</span>

The code cell for this project has an assignment statement that creates a top level Tk application named `hello_app`.  

Add statements to the code cell so the program has a label that says "Hello, World":

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/hello_gui.png"/>

The example above uses 24-point Helvetica.  You can use the default text, but feel free to experiment with other fonts, font sizes, colors, *etc*.

The goal for this part of the project is just to make sure you have all the necessary software installed and can create a simple application.  There is no documentation requirement for this part.

##### <span style="color:red">Code:</span>

<!-- grade_id: hello, points: 4, grade: true, solution: true -->
```python
hello_app = tk.Tk()
hello = tk.Label(hello_app, text = "Hello, World!")
hello.config(font = ('Papyrus', 72, 'bold'))
hello.pack()
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: hello_test_1, points: 2, grade: true, solution: false -->
```python
# The GUI should have a single widget
widgets = list(hello_app.children.values())
assert len(widgets) == 1
```

<!-- grade_id: hello_test_2, points: 2, grade: true, solution: false -->
```python
# Check to see if the widget is a Label
widgets = list(hello_app.children.values())
assert isinstance(widgets[0], tk.Label)
```

<!-- grade_id: hello_test_3, points: 2, grade: true, solution: false -->
```python
# The text in the label should start with "hello" (upper or lower case)
widgets = list(hello_app.children.values())
assert widgets[0]['text'].lower().startswith('hello')
```

### <span style="color:teal">2. &nbsp; Payment Calculator (50 points)</span>

Create a program that computes the monthly payment on a loan. The payment is a function of the loan amount, the annual interest rate, and the number of years to pay off the loan. To calculate the payment, first compute two values named $r$ and $p$, defined by these equations:

$ r = ( \textrm{rate} \; / \; 100 ) \; / \; 12 $

$ p = 12 \times \textrm{years} $

The monthly payment is then

$$ \textrm{payment} = \frac{ r \times \textrm{amount} }{ 1 - (1 + r)^{-p} } $$

For example, if you take out a 30-year loan for \$150,000 at an annual rate of 4.5% your payment would be \$760.03 per month.

The GUI for this application should have four text entry boxes, and each box should have a label to its left that explains what the box holds.  The first three boxes are for the user to enter the loan amount, the interest rate, and the number of years to repay the loan. 

The GUI should have a button labeled “compute payment”. When the user clicks this button, your program should read the values of the three parameters, use the equations shown above to calculate the monthly payment, and display the payment amount in the fourth text entry box.

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/payment_gui.png"/>

The code cell for this program has the outline for the program.  First fill in the body of the function named `monthly_payment` and make sure it computes the correct payment value given an amount, rate, and loan period.  When your write the button callback function have it call the `monthly_payment` function.

Next add statements that create the labels, boxes, and button for the application.  When you create the Entry and Button widgets, make sure you save them in global variables with the following names so the autograder tests will succeed:
* `amount`
* `rate`
* `years`
* `payment`
* `button`

##### <span style="color:red">Code:</span>

<!-- grade_id: payment, points: 15, grade: true, solution: true -->
```python
def monthly_payment(amt, pct, yrs): 
    '''
    Given the amount of money, the interest rate, and number of years on the payment, returns a monthly payment.
    
    :param amt str/float/int: The amount of money that is being borrowed.
    :param pct str/float/int: The interest rate in percentage.
    :param yrs str/int: The number of years over which the payment will be made. 
    :return: Monthly payment based on the parameters given.
    :rtype: float.
    '''
    r = (float(pct) / 100) / 12
    p = float(yrs) * 12
    m_payment = (r * float(amt)) / (1 - (1 + r) ** -p)
    return m_payment

def compute_payment_cb():
    '''
    Uses monthly_payment() to calculate the monthly payment based on the values input into the GUI. 
    Returns the monthly payment to the payment box of the GUI.
    
    :return: None.
    :rtype: None.
    '''
    # YOUR CODE HERE
    a = amount.get()
    ra = rate.get()
    y = years.get()
    x = monthly_payment(a, ra, y)
    payment.delete(0, tk.END)
    payment.insert(0, "${:.2f}".format(x))

payment_app = tk.Tk()

amount_lab = tk.Label(payment_app, text = "        Loan Amount:                                   ", font = ("Comic Sans MS", 30), fg = '#ff69b4', bg = "#5A6351")
amount_lab.grid(row = 0, column = 0, padx = 40, pady = 10)
amount = tk.Entry(payment_app)
amount.grid(row = 0, column = 1, padx = 70, pady = 20)

rate_lab = tk.Label(payment_app, text = "                                  Interest Rate:  ", font = ("Chalkduster", 30), fg = '#ffff00', bg = '#e5e500')
rate_lab.grid(row = 1, column = 0, padx = 0 , pady = 0)
rate = tk.Entry(payment_app)
rate.grid(row = 1, column = 1, padx = 20, pady = 20)

years_lab = tk.Label(payment_app, text = " Loan Period:                                                  ", font = ('Futura', 30), fg = '#00FA9A', bg = '#0000e5')
years_lab.grid(row=2, column=0, padx = 20, pady = 0)
years = tk.Entry(payment_app)
years.grid(row=2,column=1, padx = 25, pady = 50)

payment_lab = tk.Label(payment_app, text = "               Payment:               ", font = ('Papyrus', 25, 'underline'), fg = '#D2691E', bg = '#83F52C')
payment_lab.grid(row=3, column=0, padx = 60, pady = 60)
payment = tk.Entry(payment_app)
payment.grid(row=3, column=1, padx = 40, pady = 10)

button = tk.Button(payment_app, text = " Compute Payment  ", font = ('Times New Roman', 32, 'bold'), fg = '#ff0000', bg = '#000000', command = compute_payment_cb)
button.grid(row=4, column = 0, columnspan = 2, pady = 10)
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: payment_test_1, points: 5, grade: true, solution: false -->
```python
# Test the callback function to make sure it is computing payments accurately
assert round(monthly_payment(20000, 3.0, 6), 2) == 303.87
```

<!-- grade_id: payment_test_2, points: 10, grade: true, solution: false -->
```python
# Verify the number of widgets of each type
counts = { tk.Label: 0, tk.Entry: 0, tk.Button: 0 }

for x in payment_app.children.values():
    counts[type(x)] += 1
    
assert counts == { tk.Label: 4, tk.Entry: 4, tk.Button: 1 }
```

<!-- grade_id: payment_test_3, points: 10, grade: true, solution: false -->
```python
# Store values in the GUI, click the button, check the result
amount.insert(0, '20000')
rate.insert(0, '3.0')
years.insert(0, '6')
button.invoke()

assert payment.get() == '$303.87'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
monthly_payment has three parameters: amt (the amount of money that is getting borrowed), pct (The interest rate expressed in percentage form, i.e. a value of 5 will be interpreted at 5% and a value of .05 will be used in the equations), and yrs (the number of years over which the loan will be paid. 

The function first defines r and p as they were in the equation given for ease of reading. Then m_payment is assigned to the value of the equation with the parameters plugged in. The function then returns m_payment.


compute_payment_cb first defines a, ra, and y to be the values inputted into the GUI for amount, rate, and years respectively. ra was chosen instead of r in order to avoid confusion with the variable in monthly_payment. It then defines x to be the value of monthly_payment with those variables. It then deletes the content of the payment text entry box and inserts the value of x into it. 


payment_app is the top level application that we are basing the entire GUI on. 

amount_lab, rate_lab, years_lab, and payment_lab are the labels for amount, rate, years, and payment respectively. amount is a text entry box for the amount of money being borrowed, rate is a text entry box for the interest rate, years is a text entry box for the number of years that the loan is taking place, and payment is the text entry box where the car payment is going to appear. 

button is a button with "Compute Payment" on it that will be clicked when the person wants to compute their monthly payment. It executes compute_payment_cb. 

All of these have .grid() functions on them in order to create a more aestetically pleasing GUI.



### <span style="color:teal">3. &nbsp; Canvas (40 points)</span>

Fill in the application named `canvas_app`.  The application will have just one widget, a Tk canvas with a size of $400 \times 400$ pixels. 

You also need to fill in the body of the function named `make_circle` that will draw a circle on your canvas.  The parameters are a reference to the canvas widget, three integers for the size and location of the circle, and a string with a color name.

Use the function to draw three circles on your canvas.  You can choose whatever size, location, and color you would like.

<img src="http://www.cs.uoregon.edu/Classes/16W/cis211/images/gui/three_circles.png"/>


##### <span style="color:red">Code:</span>

<!-- grade_id: canvas, points: 15, grade: true, solution: true -->
```python
def make_circle(canvas, radius, x, y, color):
    drawing.create_oval((x-radius, y-radius), (x + radius, y+radius), fill = color)

canvas_app = tk.Tk()
drawing = tk.Canvas(canvas_app, height = 400, width = 400, background = "lightgray")
drawing.pack()
make_circle(canvas_app, 20, 100, 200, '#000000')
make_circle(canvas_app, 50, 300, 300, '#FF0000')
make_circle(canvas_app, 70, 200, 200, '#FFFF00')
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: canvas_test_1, points: 5, grade: true, solution: false -->
```python
# See if the GUI has a Canvas widget
widgets = list(canvas_app.children.values())
assert len(widgets) == 1 and isinstance(widgets[0], tk.Canvas)
```

<!-- grade_id: canvas_test_2, points: 10, grade: true, solution: false -->
```python
# See if the canvas has circles on it
widgets = list(canvas_app.children.values())
for n in widgets[0].find_all():
    assert widgets[0].type(n) == 'oval'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_3, points: 10, grade: true, solution: true -->
make_circle uses tk.Canvas to create a circle with a specified radius, with a center at (x,y) with a color of the user's choosing. It uses create_oval to create edges of the circle at (x-radius, y-radius) and (x+radius, y+radius) to ensure that the (x,y) coordinate chosen is at the center. 

canvas_app is our top level application. 

drawing creates a 400 x 400 pixel canvas with a light grey background, which then is made visible by drawing.pack().
Three circles are then made: a black circle at (100, 200) with radius 20, a red circle at (300,300) with radius 50, and a yellow circle at (200,200) with radius 70.   